Crosslinking by ZapD drives the assembly of short FtsZ filaments into toroidal structures in solution

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

The major result in the manuscript is the observation of the higher order structures in a cryoET reconstruction that could be used for understanding the assembly of toroid structures. The cross-linking ability of ZapD dimers result in bending of FtsZ filaments to a constant curvature. Many such short filaments are stitched together to form a toroid like structure. The geometry of assembly of filaments - whether they form straight bundles or toroid like structures - depends on the relative concentrations of FtsZ and ZapD.

Strengths:

In addition to a clear picture of the FtsZ assembly into ring-like structures, the authors have carried out basic biochemistry and biophysical techniques to assay the GTPase activity, the kinetics of assembly, and the ZapD to FtsZ ratio.

Weaknesses:

The discussion does not provide an overall perspective that correlates the cryoET structural organisation of filaments with the biophysical data. The current version has improved in terms of addressing this weakness and clearly states the lacuna in the model proposed based on the technical limitations.

Future scope of work includes the molecular basis of curvature generation and how molecular features of FtsZ and ZapD affect the membrane binding of the higher order assembly.

Reviewer #3 (Public review):

Summary:

Previous studies have analyzed the binding of ZapD to FtsZ and provided images of negatively stained toroids and straight bundles, where FtsZ filaments are presumably crosslinked by ZapD dimers. Toroids without ZapD have also been previously formed by treating FtsZ with crowding agents. The present study is the first to apply cryoEM tomography, which can resolve the structure of the toroids in 3D. This shows a complex mixture of filaments and sheets irregularly stacked in the Z direction and spaced radially. The most important interpretation would be to distinguish FtsZ filaments from ZapD crosslinks, This is less convincing. The authors seem aware of the ambiguity: "However, we were unable to obtain detailed structural information about the ZapD connectors due to the heterogeneity and density of the toroidal structures, which showed significant variability in the conformations of the connections between the filaments in all directions." Therefore, the reader may assume that the crosslinks identified and colored red are only suggestions, and look for their own structural interpretations. But readers should also note some inconsistencies in stoichiometry and crosslinking arrangements that are detailed under "weaknesses."

Strengths.

This is the first cryoEM tomography to image toroids and straight bundles of FtsZ filaments bound to ZapD. A strength is the resolution, which. at least for the straight bundles. is sufficient to resolve the ~4.5 nm spacing of ZapD dimers attached to and projecting subunits of an FtsZ filament. Another strength is the pelleting assay to determine the stoichiometry of ZapD:FtsZ (although this also leads to weaknesses of interpretation).

Weaknesses

The stoichiometry presents some problems. Fig. S5 uses pelleting to convincingly establish the stoichiometry of ZapD:FtsZ. Although ZapD is a dimer, the concentration of ZapD is always expressed as that of its subunit monomers. Fig. S5 shows the stoichiometry of ZapD:FtsZ to be 1:1 or 2:1 at equimolar or high concentrations of ZapD. Thus at equimolar ZapD, each ZapD dimer should bridge two FtsZ's, likely forming crosslinks between filaments. At high ZapD, each FtsZ should have it's own ZapD dimer. However, this seems contradicted by later statements in Discussion and Results. (1) "At lower concentrations of ZapD, .. toroids are the most prominent structures, containing one ZapD dimer for every four to six FtsZ molecules." Shouldn't it be one ZapD dimer for every two FtsZ? (2) "at the high ZapD concentration...a ZapD dimer binds two FtsZ molecules connecting two filaments." Doesn't Fig. S5 show that each FtsZ subunit has its own ZapD dimer? And wouldn't this saturate the CTD sites with dimers and thus minimize crosslinking?

We thank the reviewer for these insightful comments. The affinity of ZapD for FtsZ is relatively low and a higher concentration of ZapD is required in solution to effectively saturate the binding sites of all FtsZ molecules forming macrostructures. It is important to clarify that the concentrations mentioned in the text refer to the amounts and ratios of protein added to the total volume of the sample, rather than the proteins actively interacting and forming bundles or macrostructures.

To differentiate, two aspects can be considered: the ratio of added protein (as mentioned in the text) and the fraction of proteins that contribute to the formation of the macrostructures. Under polymerization conditions, FtsZ-GTP recruits additional monomers to form polymers. Therefore, more FtsZ than ZapD would be involved in forming filaments and bundles. Our results support this hypothesis and show that a higher amount of ZapD is required in the sample to pellet with FtsZ bundles.

We propose that starting with the same initial concentration of FtsZ and ZapD in solution, only a small fraction of ZapD will bind to the structures, favoring the formation of toroidal structures despite the initial 1:1 ratio of proteins added to the sample. When considering a higher FtsZ:ZapD ratio (1:6), the increased amount of ZapD in solution would facilitate the saturation of all FtsZ binding sites, consistent with the observation of straight bundles. Analytical sedimentation velocity data further supported this finding, indicating a binding ratio of approximately 0.3-0.4, suggesting that one ZapD dimer binds for every 4-6 FtsZ monomers. The binding ratio indicates that two FtsZ monomers will bind to a single dimer of ZapD, but this only occurs when there is a significant excess of ZapD over FtsZ in the solution mixture.

These findings align qualitatively with the relative intensities of the electrophoretic bands observed for FtsZ and ZapD in the pelleting assay with different FtsZ-ZapD mixtures, as shown in Suppl. Fig. 5 as % of FtsZ in the fractions. Without prior staining calibration of the gels, there is no simple quantitative relationship between gel band intensities after Coomassie staining and the amount of protein in a band (Darawshe et al. 1993 Anal Biochem - DOI: 10.1006/abio.1993.1581). This last point precludes a quantitative comparison between pelleting/SDS-PAGE data and analytical sedimentation measurements. For this reason, we have decided to present pelleting results as % of FtsZ in supernatant and pellet to avoid overestimations.

A major weakness is the interpretation of the cryoEM tomograms, specifically distinguishing ZapD from FtsZ. The distinction of crosslinks seems based primarily on structure: long continuous filaments (which often appear as sheets) are FtsZ, and small masses between filaments are ZapD. The density of crosslinks seems to vary substantially over different parts of the figures. More important, the density of ZapD's identified and colored red seem much lower than the stoichiometry detailed above. Since the mass of the ZapD monomer is half that of FtsZ, the 1:1 stoichiometry in toroids means that 1/3 of the mass should be ZapD and 2/3 FtsZ. However, the connections identified as ZapD seem much fewer than the expected 1/3 of the mass. The authors conclude that connections run horizontally, diagonally and vertically, which implies no regularity. This seems likely, but as I would suggest that readers need to consider for themselves what they would identify as a crosslink.

The amount of ZapD in the toroids will be significantly less than one third. Although the theoretical addition of protein to the samples is at a 1:1 ratio, the actual amount of protein in the macrostructures containing ZapD is much lower, as shown by sedimentation velocity pelleting assays.

In contrast to the toroids formed at equimolar FtsZ and ZapD, thin bundles of straight filaments are assembled in excess ZapD. Here the stoichiometry is 2:1, which would mean that every FtsZ should have a bound ZapD DIMER. The segmentation of a single filament in Fig. 5e seems to agree with this, showing an FtsZ filament with spikes emanating like a picket fence, with a 4.5 nm periodicity. This is consistent with each spike being a ZapD dimer, and every FtsZ subunit along the filament having a bound ZapD dimer. But if each FtsZ has its own dimer, this would seem to eliminate crosslinking. The interpretative diagram in Fig. 6, far right, which shows almost all ZapD dimers bridging two FtsZs on opposite filaments, would be inconsistent with this 2:1 stoichiometry.

Assessing the precise stoichiometry of FtsZ and ZapD within the macrostructures is challenging. We interpret the spikes as ZapD dimers bridging two FtsZ filaments, implying a theoretical 1:1 stoichiometry in the straight bundle. However, ZapD may be enriched in certain areas, indicating that a single FtsZ monomer is binding to one side of the dimer. In contrast, the other side remains available for additional connections, resulting in a potential 2:1 stoichiometry. A combination of both scenarios is likely, although our resolution does not allow further characterization. Considering these complexities, we assume these connections represent a dimer of ZapD binding to two FtsZ monomers.

Figure 6 shows a simplified scheme illustrating how the bundles could be assembled based on the Cryo-ET data. We acknowledge the limitations of this diagram; its purpose is to depict the mesh formed by the stabilization of ZapD. We have not included interactions that do not lead to filament crosslinking, such as dimers binding to only one FtsZ filament. This focus enhances the interpretation of the scheme and the FtsZ-ZapD interaction. A sentence has been added to the caption to highlight the possibility of other interactions not considered in the scheme.

In the original review I suggested a control that might help identify the structures of ZapD in the toroids. Popp et al (Biopolymers 2009) generated FtsZ toroids that were identical in size and shape to those here, but lacking ZapD. These toroids of pure FtsZ were generated by adding 8% polyvinyl chloride, a crowding agent. The filamentous substructure of these toroids in negative stain seemed very similar to that of the ZapD toroids here. CryoET of these toroids lacking ZapD might have been helpful in confirming the identification of ZapD crosslinks in the present toroids. However, the authors declined to explore this control.

The mechanisms by which methylcellulose (MC) promotes the assembly of FtsZ macrostructures reported by Popp et al. involve more than simple excluded volume effects, as the low concentration of MC (less than 1 mg/ml) falls below the typical crowding regime. The latter suggests the existence of poorly characterized additional interactions between MC and FtsZ. These complexities preclude the use of FtsZ polymers formed in the presence of MC as a true control for the FtsZ toroidal structures reported here.

Finally, it should be noted that the CTD binding sites for ZapD should be on the outside of curved filaments, the side facing the membrane in the cell. All bound ZapD should project radially outward, and if it contacted the back side of the next filament, it should not bind (because the CTD is on the front side). The diagram second to right in Fig. 6 seems to incorporate this abortive contact.

The role of the flexible linker and its biological implications are still under debate in the field. The flexible linker allows ZapD-driven connections to be made in different directions. While these implications are not the primary focus of our manuscript, the flexible linker could allow connections between filaments in different orientations.

Reviewer #1 (Recommendations for the authors):

Most of the concerns which I had raised in the earlier version have been taken care of, as detailed in the response.

A few minor points, mostly related to re-phrasing are listed below:

Page 2: line 21: The use of the term 'C-terminal domain' for the C-terminal unstructured region of FtsZ is confusing. The term C-terminal domain or CTD for FtsZ is commonly used to describe part of the globular domain, while C-terminal tail or CCTP will be a more apt usage for all the instances in this manuscript.

We refer to the C-terminal domain as the carboxy-terminal region of the protein. This domain includes the C-terminal linker (CTL), which varies in length between species, followed by a conserved 11-residue sequence (CTC) and shorter, variable C-terminal sequences (CTV). We used the term "C-terminal domain" primarily to improve the readability of the manuscript, but we appreciate the reviewer's feedback. We have now adopted the term "CCTP" instead of "C-terminal domain" to improve the clarity of our manuscript.

On a related note, the schematic in Fig 1 shows the interaction with CCTP rather than the C-terminal domain of the globular FtsZ. Please provide an explanation.

We refer to the unstructured C-terminal domain of FtsZ as the C-terminal tail. To avoid confusion, we have introduced the term CCTP in this manuscript.

Supple Fig 2: "The FCS analysis demonstrated an increasing diffusion time of ZapD along with the FtsZ concentration as result of higher proportion of ZapD bound to FtsZ.

The increased diffusion time need not be interpreted as increased ZapD bound, it could also mean that FtsZ could polymerise in the presence of increasing ZapD, was this possibility ruled out? Including a comment on this aspect will be useful.

In these experiments, we monitored fluorescently labeled ZapD. Due to their interaction, we found that its diffusion time increased at high FtsZ concentrations. The data presented in Supplementary Figure 2 shows ZapD in the presence of FtsZ-GDP (i.e. under non-polymerization conditions).

Was it possible to get a molecular weight estimate based on the diffusion time?

It is possible to estimate hydrodynamic volumes using the Stokes-Einstein equation if the diffusion coefficient of the diffusing particles is known, assuming that the particles are small and spherical. A molecular weight can then be estimated using a standard density of 1.35 g/cm3 (Fisher et all. Protein science 2009 DOI: 10.1110/ps.04688204). This estimate is heavily dependent on the shape of the diffusing particle, as we assume that our protein of interest here is far from a spherical shape due to the interaction through the flexible linker, the hydrodynamic volumes are overestimated. This overestimation then leads to a further overestimation of the molecular weight. In addition, for a more accurate estimation of the sizes and thus molecular weights for proteins, a modified model of the Stokes-Einstein equation is required (Tyn and Gusek Biotechnology and Bioengineering DOI: 10/1002/bit.260350402), where additional information about the shape of the diffusing particle is estimated by measuring the radius of gyration of the particle. These calculations are complex and beyond the scope of our manuscript.

Supple Fig 4:

Does FtsZ GTPase activity (without ZapD) also vary with KCl concentrations? It will be useful to comment on this in Supplementary Figure 4.

Yes, it has been previously reported that moderate concentration of KCl is optimal for FtsZ GTPase activity. We added a comment to the caption.

Page 6, line 42: short filament segments arranged nearly 'parallel' to each other Since FtsZ filaments are polar, it is better to rephrase as 'parallel or antiparallel'.

Corrected.

Page 7, line 41: cross linking of short 'FtsZ' filaments and not ZapD?

It was a typo. Corrected

Page 8: delete 'from above' in the title?

Corrected

The use of the phrases such as 'cross linking from the top'; 'binds to FtsZ from above' is vague. (Figure 5b legend; discussion page 10, line 18; page 8, line 26; page 12, line 27). Similarly labelling on a schematic figure on the use of vertical, diagonal/lateral will be useful for the readers.

We thank the reviewer for the suggestions to improve the understanding of our data. We have simplified them by renaming these interactions as vertical.

Page 13, lines 6 -10

Rather than an orientation of top or from the side, just the presence of multiple crosslinks along coaxial filaments suffices for a straight bundle. The average spacing will be more uniform in such a straight bundle compared to a toroid where there might be regions without ZapD. I do not find the data on an upward orientation convincing. ZapD binding need not be above to have the C-terminal ends of FtsZ pointing towards the membrane. On the other hand, having ZapD bind above is likely to occlude membrane binding of FtsZ?

The flexibility of the FtsZ linker suggests that ZapD can bind filaments oriented in different directions. In a cellular environment, FtsZ molecules interact with other division proteins that compete with ZapD for binding sites. This competition could prevent the membrane from occluding and instead create binding sites between the filaments, stabilizing them.

Page 11, lines 32 - 34: Please rephrase the sentence, with focus on the main point to be conveyed. Do the authors want to say that the 'Same molecule contributes to variability in spacing based on the number of connections formed.'

Thank you for your comment. We have rephrased the sentence for clarity.

Page 11: paragraphs 1,2, and 3 appears to convey similar, related ideas and are redundant. Could these be shortened further into one paragraph highlighting how the ratio leads to differences in higher order FtsZ organisation?

These paragraphs discuss different ideas, and it is better to keep them separate.

In the response to reviewers, page 19, point 5 (iii), it is given that 5000 FtsZ molecules correspond to 2/3rd of the total, while in the manuscript text, it is given as one-third. Please correct the response text/manuscript text accordingly. The numbers in the cited reference appears to suggest 1/3rd.

Yes, it was 1/3rd. Thanks for pointing that out.

Fig 1b. Y-axis: Absorbance spelling has a typo.

Page 14, line 11: Healthcare ('h' missing)

Page 14, line 15: HCl, KCl (L should be in small letter)

Page15, line 18: 43 - 48K rpm (not Krpm)

Supple Fig 1 legend: line 5: 's' missing for species

Corrected.

Changes from the earlier version to the revised version

Abstract

Abstract

Bacterial cell division relies on the Z ring, a cytoskeletal structure that acts as a scaffold for the assembly of the divisome. To date, the detailed mechanisms underlying the assembly and stabilization of the Z ring remain elusive. This study highlights the role of the FtsZ-associated protein (Zap) ZapD in the assembly and stabilization of Z-ring-like structures via filament crosslinking. Using cryo-electron tomography and biochemical analysis, we show that, at equimolar concentrations of ZapD and FtsZ, ZapD induces the formation of toroidal structures composed of short, curved FtsZ filaments that are crosslinked vertically, but also laterally and diagonally. At higher concentrations of ZapD, regularly spaced ZapD dimers crosslink FtsZ filaments from above, resulting in the formation of straight bundles. Despite the simplicity of this reconstituted system, these findings provide valuable insights into the structural organization and stabilization of the Z ring by Zap proteins in bacterial cells, revealing the key role of optimal crosslinking density and geometry in enabling filament curvature and ring formation.

Introduction

Cell division in bacteria is a well-orchestrated process mediated by a multiprotein complex called the divisome. Their components interact reversibly to form a division ring, which is essential for cytokinesis1–3. At the core of this process in most bacteria is FtsZ, a tubulin homolog that polymerizes in the presence of GTP beneath the inner membrane at the future division site, providing a scaffold for recruiting other division proteins4–8. FtsZ undergoes treadmilling due to polymerization-dependent GTP hydrolysis, allowing the ring to exhibit its dynamic behavior9. Depending on environmental conditions, FtsZ filaments can assemble into various structures, including bundles, sheets, and toroids10–14. Membrane-tethered FtsZ assemblies are arranged as loosely structured spirals in the midcell region, guided by spatial regulators. These structures eventually coalesce to form a cohesive Z ring15.

Earlier version: The molecular components involved in cell division must be in the right place at the right time during the cell cycle. Several factors in this spatiotemporal regulation modulate FtsZ assembly. In Escherichia coli, FtsA and ZipA tether FtsZ to the membrane16–18 and, together with a set of FtsZ-associated proteins (Zaps), stabilize the ring19–22. Conversely, negative modulators inhibit FtsZ assembly at other sites, namely Min proteins at the polar regions23 and SlmA on the nucleoid24. Most of these modulators interact with FtsZ through its carboxy-terminal end, which modulates division assembly as a central hub25,26. ZapD is the only Zap protein known to crosslink FtsZ by binding its C-terminal domain, suggesting a critical Z ring structure stabilizing function19,27–30 (Fig. 1a).

Revised version: The molecular components involved in cell division must be in the right place at the right time during the cell cycle. Several factors in this spatiotemporal regulation modulate FtsZ assembly. In Escherichia coli, FtsA and ZipA tether FtsZ to the membrane16–18 and, together with a set of FtsZ-associated proteins (Zaps), stabilize the ring19–22. Conversely, negative modulators inhibit FtsZ assembly at other sites, namely Min proteins at the polar regions23 and SlmA on the nucleoid24. Most of these modulators interact with FtsZ through its carboxy-terminal end (CCTP), which modulates division assembly as a central hub25,26. ZapD is the only Zap protein known to crosslink FtsZ by binding CCTP, suggesting a critical Z ring structure stabilizing function19,27–30 (Fig. 1a).

ZapD binds FtsZ and promotes filament bundling.

Earlier version: a. Scheme of the FtsZ protein and its interaction with ZapD. E. coli FtsZ (blue) monomers in solution oligomerize depending on the buffer conditions. Upon GTP binding, FtsZ homopolymerizes directionally, assembling single stranded filaments. ZapD is a robust dimer (magenta) that interacts directly with FtsZ, crosslinking filaments by promoting lateral interactions between them. Although the molecular mechanism is still unclear, the current hypothesis of interaction assumes dimers of ZapD crosslinking two FtsZ filaments through the C-terminal region of FtsZ, expecting at around 1:1 (FtsZ:ZapD) molar ratio in an homogeneous bundle (1 dimer of ZapD connected to 2 monomers of FtsZ). According to this model, the orientation of the FtsZ filaments could be parallel or antiparallel, allowing the growth and treadmilling of the filaments. However, the mechanism of assembly of dynamic high-order structures is still unknown.

Revised version: a. Scheme of the FtsZ protein and its interaction with ZapD. E. coli FtsZ (blue) monomers in solution oligomerize depending on the buffer conditions. Upon GTP binding, FtsZ homopolymerizes directionally, assembling single stranded filaments. ZapD is a robust dimer (magenta) that interacts directly with FtsZ, crosslinking filaments by promoting lateral interactions between them. Although the molecular mechanism is still unclear, the current hypothesis of interaction assumes dimers of ZapD crosslinking two FtsZ filaments through the CCTP of FtsZ, expecting at around 1:1 (FtsZ:ZapD) molar ratio in a homogeneous bundle (1 dimer of ZapD connected to 2 monomers of FtsZ). According to this model, the orientation of the FtsZ filaments could be parallel or antiparallel, allowing the growth and treadmilling of the filaments. However, the mechanism of assembly of dynamic high-order structures is still unknown.

b. Turbidity assays measuring the absorbance at 350 nm of samples containing 5 µM FtsZ and increasing concentrations of ZapD. The turbidity of the sample was measured 5 minutes after the addition of 1 mM GTP at working buffer (50 mM KCl, 50 mM Tris-Cl, 5 mM MgCl2 pH 7). FtsZ polymers do not show a significant turbidity at this wavelength; therefore, the signal at 350 nm corresponds to the presence of large FtsZ macrostructures and bundles. The mean value and SD of >3 independent replicates are plotted in the graph.

c. GTPase activity of FtsZ after addition of 1 mM GTP in presence of increasing concentrations of ZapD at working conditions (50 mM KCl, 50 mM Tris-Cl, 5 mM MgCl2 pH 7). The mean value and SD plotted in the graph are the result of 3 independent replicates. The GTPase activity was measured as a result of the Pi released from GTP consumption. The units are Mol GTP consumed per Mol FtsZ per min.

To date, a high-resolution structure of the Z ring remains unresolved. Most evidence suggests that FtsZ filaments are arranged in short patches along the equatorial circle, held together by lateral interactions and interactions with Zap proteins2,7,31–33. The dynamics of treadmilling and lateral interactions between FtsZ filaments, as well as crosslinking by FtsZ-associated proteins (Zaps), are thought to play a crucial role in the condensation of the Z ring34. Specifically, the crosslinking of filaments by Zaps, particularly ZapD, is vital for maintaining the neighboring organization of the filaments into a ring structure. This organization may also be critical for the functionality of the divisome and the mechanism of force generation during cytokinesis1,2,7,35–38.

Zap proteins play overlapping roles in stabilizing the Z ring during cell division. Although they are non-essential for division, their absence leads to less compact Z rings, and the know-out of two or more Zap encoding genes simultaneously results in cell elongation19,39,40. Notably, these different Zap proteins do not share any sequence homology and differ in their structure and filament crosslinking mechanism20,27,41. Additionally, Zap proteins may actively remove FtsZ from the septum during the final stages of constriction42. They also facilitate the positioning of the Z ring at the replication terminus of the chromosome by interacting with MatP, which forms a scaffolding anchor known as Ter-linkage43.

Earlier version: ZapD is a symmetrical dimer consisting of an α-helical domain and a β-sheet domain containing a positively charged binding pocket required for crosslinking and bundling FtsZ filaments30,41. ZapD binds to the C-terminal region of FtsZ through electrostatic interactions30. ZapD crosslinks FtsZ filaments, promoting bundling and significantly reducing the GTPase activity of FtsZ19. A model of how ZapD crosslinks neighboring FtsZ filaments has been proposed based on the structural analysis of ZapD-FtsZ interactions41. In this model, the ZapD dimer connects two FtsZ filaments through the C-terminal domain, which could organize them in a parallel or antiparallel orientation due to the flexibility provided by the flexible linker (Fig. 1a). However, the experimental validation of this model is still missing.

Revised version: ZapD is a symmetrical dimer consisting of an α-helical domain and a β-sheet domain containing a positively charged binding pocket required for crosslinking and bundling FtsZ filaments30,41. ZapD binds to the CCTP of FtsZ through electrostatic interactions30. ZapD crosslinks FtsZ filaments, promoting bundling and significantly reducing the GTPase activity of FtsZ19. A model of how ZapD crosslinks neighboring FtsZ filaments has been proposed based on the structural analysis of ZapD-FtsZ interactions41. In this model, the ZapD dimer connects two FtsZ filaments through the CCTP, which could organize them in a parallel or antiparallel orientation due to the flexibility provided by the flexible linker (Fig. 1a). However, the experimental validation of this model is still missing.

Furthermore, the link between FtsZ bundling, promoted by ZapD, and the large-scale organization of FtsZ filaments remains unresolved. Therefore, investigating how ZapD crosslinks FtsZ filaments may shed light on the molecular mechanisms underlying Z ring assembly. Recently, Gong et al.44 demonstrated that although the FtsZ crosslinking proteins are not individually essential, they collectively play a critical role in the condensation and stability of the Z ring in different organisms. However, the specific mechanisms driving this condensation remain unclear, highlighting the complexity of the bacterial cell division process.

In this study, we used cryo-electron microscopy (cryo-EM) and cryo-electron tomography (cryo-ET), together with biochemical and biophysical assays, to investigate the structural organization of FtsZ polymers in the presence of ZapD. This integrated approach revealed that, at equimolar concentrations of ZapD and FtsZ, ZapD facilitates the assembly of FtsZ filaments into toroidal polymers in solution. This observation is consistent with the low resolution structure of the Z-ring obtained in vivo7,45–47. Our results allowed us to propose a molecular mechanism for toroid formation, providing valuable insights into how crosslinking of division proteins drives the stabilization of the Z-ring.

Results

ZapD dimer interacts with FtsZ-GDP oligomers

Earlier version: Previous structural studies have identified a direct interaction between ZapD and FtsZ through the C-terminal domain or central hub of FtsZ, involving charged residues in the interaction28,30. To independently verify these findings, we investigated the interaction between FtsZ-GDP and ZapD using analytical ultracentrifugation (AUC) (see Supplementary Fig. 1). First, we confirmed that ZapD forms stable dimers, which is consistent with previous studies19,28,41 (Supplementary Fig. 1a (I) and b). In contrast, unpolymerized FtsZ-GDP primarily self-assembles from monomers into dimers and other oligomeric species, albeit in small proportions, as expected48 (Supplementary Fig. 1a (II)).

Revised version: Previous structural studies have identified a direct interaction between ZapD and FtsZ through the CCTP or central hub of FtsZ, involving charged residues in the interaction28,30. To independently verify these findings, we investigated the interaction between FtsZ-GDP and ZapD using analytical ultracentrifugation (AUC) (see Supplementary Fig. 1). First, we confirmed that ZapD forms stable dimers, which is consistent with previous studies19,28,41 (Supplementary Fig. 1a (I) and b). In contrast, unpolymerized FtsZ-GDP primarily self- assembles from monomers into dimers and other oligomeric species, albeit in small proportions, as expected48 (Supplementary Fig. 1a (II)).

Sedimentation velocity analysis of mixtures of the two proteins revealed the presence of two predominant molecular species of ZapD:FtsZ complexes in solution. These complexes are compatible with ZapD dimers bound to one or two FtsZ monomers, corresponding to ZapD:FtsZ stoichiometries of 2:1 and 1:1, respectively (Supplementary Fig. 1a (III-IV)). This observation is consistent with the proposed interaction model41. Furthermore, we confirmed the interaction of both proteins using fluorescence anisotropy and fluorescence correlation spectroscopy (Supplementary Fig. 2a and b). In these experiments, an increase in fluorescence anisotropy or diffusion time suggested the formation of larger particles due to protein-protein interactions.

ZapD binding promotes the bundling of FtsZ-GTP polymers

ZapD is a crosslinker for FtsZ filaments, promoting their bundling in solution19. However, the precise molecular mechanism by which ZapD crosslinks the filaments remains unclear. To investigate the biochemical features underlying this bundling process, we examined how the formation of large FtsZ bundles or higher-order FtsZ polymer structures depends on the concentration of ZapD. For this analysis, we used turbidity measurements at 350 nm. When testing GTP-FtsZ filaments at 5 or 10 μM concentrations, we found negligible turbidity, similar to that of ZapD alone at concentrations ranging from 1 to 40 μM. However, when ZapD was added, the turbidity of the FtsZ polymers increased, reaching a maximum at approximately 5 μM of ZapD (with a ZapD-FtsZ molar ratio of around 1). This turbidity level did not change significantly even at higher ZapD concentrations. (Fig. 1b and Supplementary Fig. 3).

In the presence of ZapD, FtsZ rapidly forms higher order polymers upon addition of GTP, as indicated by turbidity assays. In contrast, the formation of single- or double-stranded FtsZ filaments without ZapD does not result in a significant increase in turbidity. Compared to FtsZ filaments observed in vitro, FtsZ bundles with ZapD exhibited reduced GTPase activity, reaching 20% at equimolar or higher concentrations of ZapD (Fig. 1c, Supplementary Fig. 4), in agreement with previous studies19. The macrostructures formed by FtsZ in the presence of ZapD were disassembled more slowly than the FtsZ filaments upon GTP consumption (Supplementary Fig. 3). Replenishment of GTP allowed for a partial recovery of the turbidity signal, suggesting that the FtsZ bundles can be rapidly reassembled (Supplementary Fig. 3e). Previous studies have shown that high concentrations of macromolecular crowders (such as Ficoll or dextran) promote the formation of dynamic FtsZ polymer networks13, observations that are consistent with our results. In these cases, the GTPase activity of FtsZ was significantly reduced compared to that of FtsZ filaments, leading to a decrease in GTPase turnover. Similar mechanisms may apply to assembly reactions involving ZapD19,28.

Earlier version: To understand the relationship between the concentration dependence of the turbidity signal – associated with the formation of FtsZ bundles or higher order structures driven by ZapD (see Fig. 1b) – and the amount of ZapD bound to the FtsZ binding sites, we carried out analytical centrifugation sedimentation velocity assays, previously used to characterize the binding of ZapD to GDP-FtsZ. This method allowed us to separate ZapD-FtsZ bundles from free ZapD in solution and calculate the concentration of unbound ZapD from the interference signal of the slowly sedimenting protein behind the rapidly sedimenting heavy ZapD-FtsZ complexes. Using the measured concentration of free ZapD together with known input concentrations of ZapD and FtsZ, we were able to calculate the concentration of bound ZapD and the binding density of ZapD in the bundles (moles of bound ZapD dimer per mole of FtsZ monomer) at different total concentrations of ZapD. Our results showed that in the 5-10 μM ZapD concentration range, the bundles exhibited an estimated ZapD binding density of approximately 0.2, meaning that one ZapD dimer is associated with every four to six FtsZ monomers. In contrast, at higher ZapD concentrations (30-40μM), the binding density increased to 0.5, indicating one ZapD dimer for every two FtsZ monomers. These results are consistent with parallel co-sedimentation assays and additional analysis by SDS-PAGE, demonstrating an enrichment of ZapD in the pellet containing the FtsZ bundles at elevated ZapD concentrations (see Supplementary Fig. 5).

Revised version: To understand the relationship between the concentration dependence of the turbidity signal - associated with the formation of FtsZ bundles or higher order structures driven by ZapD (see Fig. 1b) - and the amount of ZapD bound to the FtsZ binding sites, we carried out analytical centrifugation sedimentation velocity assays, previously used to characterize the binding of ZapD to GDP-FtsZ. This method allowed us to separate ZapD-FtsZ bundles from free ZapD in solution and calculate the concentration of unbound ZapD from the interference signal of the slowly sedimenting protein behind the rapidly sedimenting heavy ZapD-FtsZ complexes. Using the measured concentration of free ZapD together with known input concentrations of ZapD and FtsZ, we were able to calculate the concentration of bound ZapD and the binding stoichiometry of ZapD in the bundles (moles of bound ZapD in the bundles per mole of FtsZ in monomer molar units of both proteins) at different total concentrations of ZapD (see Supplementary Fig. 5). Our results showed that in the 5-10 μM ZapD concentration range, the bundles exhibited an estimated ZapD binding stoichiometry of approximately 0.3-0.4, meaning that one ZapD dimer is associated with every four to six FtsZ monomers. In contrast, at higher ZapD concentrations (30-40μM), the binding stoichiometry increased to 1.1 ± 0.2, indicating one ZapD dimer for every two FtsZ monomers. These results are qualitatively consistent with parallel co-sedimentation assays and additional analysis by SDS-PAGE, showing an increase in the relative intensity of ZapD in the pellet containing the FtsZ bundles at elevated ZapD concentrations in the mixture (see Supplementary Fig. 5).

Earlier version: These results suggest that the ZapD binding density in the polymers increases with higher ZapD concentrations. We then carried out cryo-electron microscopy (cryo-EM) and cryo-electron tomography (cryo-ET) studies to investigate whether this increase in binding density affects the structural organization of the bundles formed.

Revised version: These results suggest that the ZapD binding stoichiometry in the polymers increases with higher ZapD concentrations. We then carried out cryo-electron microscopy (cryo-EM) and cryo-electron tomography (cryo-ET) studies to investigate whether this increase in binding density affects the structural organization of the bundles formed.

ZapD facilitates the formation of higher order FtsZ-GTP toroidal structures

Previous studies have visualized bundles with similar features using negative-stain transmission electron microscopy19,30,41. Here, cryo-EM allowed us to avoid any staining artifact or structural distortion (due to adsorption or flattening upon grid drying49) and to resolve single filaments. In the presence of GTP, FtsZ polymerized into thin filaments of variable length and curvature (Fig. 2a). Upon addition of equimolar amounts of ZapD, corresponding to the slightly substoichiometric ZapD binding densities described in the previous section, FtsZ filaments assembled into bundles that predominantly formed circular toroidal structures (Fig. 2b, c and Supplementary Fig. 6a), making them a compelling study target. A close-up view of the FtsZ toroid suggests a partially ordered arrangement of the filaments in the toroidal structure (Fig. 2c).

ZapD promotes the formation of FtsZ toroids.

a. Cryo-EM micrographs of FtsZ filaments (FtsZ-GTP form) (left) and ZapD protein (right) at 10 µM under working conditions (50 mM KCl, 50 mM Tris-Cl, 5 mM MgCl2 pH 7). Scale bars are 100 nm

b. Cryo-EM images of FtsZ (10 µM) in the presence of equimolar concentrations of ZapD (ratio 1:1) and 1 mM GTP in working conditions. Cryo-EM grids were plunge frozen 2 min after GTP addition to favor the assembly of FtsZ and ZapD structures. The proteins were mixed before the polymerization was triggered by GTP addition. Scale bars are 250 nm.

c. Micrograph of an individual FtsZ toroid found under the same conditions as in (b). Close-up view of an area within the toroid is framed by a dotted black line, revealing the large amount of FtsZ filaments that form its structure.

d. Distribution of the outer diameter of the FtsZ toroid. Each toroid was measured perpendicularly in the shortest and longest axis to ensure the reliability of the measurement. The mean value and standard deviation are shown in the graph.

e. Distribution of toroidal thickness. It was measured as the result of the difference between the outer and inner diameter of each toroid. The mean value and standard deviation are shown in the graph.

Quantitative analysis of the toroidal structure revealed a conserved organization, with an outer diameter on the order of the bacterial cell size (502 ± 55 nm) (Fig. 2d), a typical thickness of 127 ± 25 nm (Fig. 2e) and an average height of 93 ± 15 nm (Supplementary Fig. 6b). By measuring the shortest and longest axes, we determined that the circularity of the structure was 0.92 ± 0.1 and 0.85 ± 0.1 for the outer and inner diameters, respectively (Supplementary Fig. 6c). This conserved toroidal structure could result from the intrinsic curvature of the FtsZ filaments10,31,50 stabilized by ZapD binding. The dimension of a ZapD dimer is ∼7 nm along its longest axis. Huecas et al.31 estimated an interfilament distance of ∼6.5-6.7 nm for toroids of Bacillus subtilis FtsZ. These authors also observed a difference in this distance as a function of the linker, suggesting that linker length modulates FtsZ-FtsZ interactions. We observed a similar spacing for double filaments (5.9 ± 0.8 nm) and a longer spacing in the presence of ZapD (7.88 ± 2.1 nm). Previous studies with ZapD have suggested that distances of 6-12 nm are possible based on the protein structure41. Longer linkers may also provide additional freedom to spread the filaments further apart and allow a greater degree of variability in the connections made by ZapD.

Previous studies have shown different FtsZ structures at different concentrations and buffer conditions. FtsZ filaments are flexible and can generate different curvatures ranging from mini rings of ∼24 nm to intermediate circular filaments of ∼300 nm or toroids of ∼500 nm in diameter (reviewed in50,51). It is reasonable to assume that FtsZ filaments can accommodate the toroidal shape promoted by ZapD crosslinking.

Our data showed that ZapD dimers crosslink adjacent FtsZ filaments into bundles and toroids. These toroids are remarkably regular in size and similar to the bacterial diameter7,45,46, suggesting a conserved intrinsic curvature of the FtsZ filaments across a range of ZapD bonds50.

3D structure of ZapD-mediated FtsZ toroids revealed by cryo-ET

Visualization of toroidal FtsZ structures using cryo-EM revealed a complex arrangement of FtsZ filaments within the toroidal shape (Fig. 2c). To gain a more detailed understanding of their three-dimensional (3D) organization, we used cryo-ET. Our initial analysis focused on ZapD-mediated toroidal FtsZ structures at equimolar concentrations of ZapD and FtsZ (Fig. 3 and Supplementary Fig. 7a).

3D structure of FtsZ toroids revealed by cryo-ET.

a. Representative tomographic slice of an FtsZ toroid resulting from the interaction of FtsZ with ZapD. The image is the average of five 0.86 nm thick tomographic slices (total thickness of 4.31 nm) of the reconstructed tomogram around the equatorial plane of a single FtsZ toroid. The concentrations of FtsZ and ZapD were 10 µM and 1 mM of GTP was added to initiate polymerization under working conditions (50 mM KCl, 50 mM Tris-Cl, 5 mM MgCl2 pH 7).

b. Close-up views of the toroidal structure show the alignment of the FtsZ filaments forming the toroid. Red arrows indicate the presence of connections between filaments.

c. The tomographic slice in the XZ plane (left) shows the cross-section corresponding to the area marked by the white dotted line in b. This image is the average of nine tomographic slices (total thickness of 7.74 nm) from the denoised tomogram. The isosurface of the cross-section (right) shows the vertical alignment and stacking of the FtsZ filaments within the toroid. This suggests that the interaction between FtsZ filaments and ZapD is mainly along the Z direction. FtsZ filaments are represented in blue.

d. Isosurface of the FtsZ toroid shown in a. It was extracted from the reconstruction of the denoised tomographic volume and positioned in different views to facilitate its visualization: (top) front view, (middle) side view and (bottom) lateral view. The toroid has a diameter of ∼552 nm and a height of ∼92 nm.

Earlier version: e. A close-up view of the segmented toroidal structure. It shows the complex internal organization of filaments assembling the toroid. It corresponds to a zone within the toroid shown in b on the right. Close-up views of the isosurface show different connections between filaments. The segmentation shown has a width of 135.9 nm x 101.48 nm and a height of 63.64 nm.

Revised version: e. A close-up view of the segmented toroidal structure. It shows the complex internal organization of filaments assembling the toroid. It corresponds to a zone within the toroid shown in b on the right. Close-up views of the isosurface show different connections between filaments. The segmentation shown has a width of 136 nm x 101 nm and a height of 64 nm.

Earlier version: Cryo-ET revealed the dense packing of the toroidal structures, with numerous densities connecting the FtsZ filaments laterally or diagonally (Fig. 3a-c, Supplementary Figs. 7a and 8a). Zoomed-in views of the toroids in the XY plane showed that the toroid consists of relatively short filament segments arranged nearly parallel to each other (Fig. 3b and Supplementary Fig. 8a). In contrast, cross-sectional views of the toroids revealed elongated structures rather than simple filament cross-sections along the Z-axis (Fig. 3c left, Supplementary Fig. 9a).

Revised version: Cryo-ET revealed the dense packing of the toroidal structures, with numerous densities connecting the FtsZ filaments laterally or diagonally (Fig. 3a-c, Supplementary Figs. 7a and 8a). Zoomed-in views of the toroids in the XY plane showed that the toroid consists of relatively short filament segments arranged nearly parallel or antiparallel to each other (Fig. 3b and Supplementary Fig. 8a). In contrast, cross-sectional views of the toroids revealed elongated structures rather than simple filament cross-sections along the Z-axis (Fig. 3c left, Supplementary Fig. 9a).

We then extracted the toroid isosurface from the tomograms to visualize its 3D structure (Fig. 3c right-e, Supplementary Fig. 8b-e and Supplementary video 1). The isosurface confirmed the presence of extended structures along the Z-axis, well beyond the elongation expected from the missing wedge effect for single FtsZ filaments (for comparison, see Supplementary Fig. 10). The vertically extended structures appeared to correspond to filaments that were connected or decorated by additional densities along the Z-axis (Supplementary Fig. 9b). Importantly, these densities were only observed in the presence of ZapD (Supplementary Fig. 10b), suggesting that they represent ZapD connections (Fig. 3e and Supplementary Figs. 8e and 9b). We note that the resolution of the data is not sufficient to precisely resolve ZapD proteins from the FtsZ filaments in the Z-axis.

Earlier version: These results suggest that the toroids are constructed and stabilized by interactions between ZapD and FtsZ, which are mainly formed along the Z-axis but also laterally and diagonally.

Revised version: These results suggest that the toroids are constructed and stabilized by interactions between ZapD and FtsZ, which are mainly formed along the Z-axis but also vertically and diagonally.

ZapD plasticity is essential for toroid shape and stabilization

Next, we manually labeled the connecting densities in the toroid isosurfaces to analyze their arrangement and connectivity with the FtsZ filaments (Fig. 4a; Supplementary Movie 2). The high density of the toroids and the wide variety of conformations of these densities prevented the use of subtomogram averaging to resolve their structure and spatial arrangement within the toroids. Most connections exhibited a characteristic bi-spherical shape between the filaments, reminiscent of ZapD dimers (Fig. 4b-d). We observed lateral connections between two parallel FtsZ filaments at the same height within the toroid (Fig. 4b). In addition, FtsZ filaments at different heights can connect vertically and diagonally through ZapD proteins acting as bridging units, forming a complex 3D mesh (Fig. 4c, d and Supplementary Movie 2). We also identified potential ZapD proteins that decorate individual FtsZ filaments and can link them to other nearby filaments (Fig. 4b). Furthermore, some filaments exhibited multiple crosslinks, leading to stronger attachments between them (Fig. 4d). Estimation of the precise number of ZapD molecules per FtsZ or the number of bonds per filament is challenging due to their inherently heterogeneous organization. However, we could observe a high number of connections stabilizing the toroidal structure. The short filament length, which results in gaps between adjacent filaments, does not correlate with a higher number of ZapD connections (Supplementary Fig. 11), indicating that ZapD is able to crosslink filaments in all directions without causing filament breakage. This could play an important mechanistic role in the functionality of the FtsZ macrostructures.

FtsZ filaments are connected by putative ZapD crosslinkers to assemble the toroidal structure.

Earlier version: a. Top (left, top), side (left, bottom) and lateral (right) views of the isosurface from a region within the toroidal structure shown in Fig. 3a. The FtsZ filaments are colored in blue, while filament connections or putative ZapD proteins are labelled in magenta to facilitate interpretation of the results. Other putative ZapD proteins decorating the FtsZ filaments were not labelled in magenta because they were not forming any clear linkage between the filaments. The segmentation shown has a width of 73.1 nm x 101.48 nm and a height of 63.84 nm.

Revised version: a. Top (left, top), side (left, bottom) and lateral (right) views of the isosurface from a region within the toroidal structure shown in Fig. 3a. The FtsZ filaments are colored in blue, while filament connections or putative ZapD proteins are labelled in magenta to facilitate interpretation of the results. Other putative ZapD proteins decorating the FtsZ filaments were not labelled in magenta because they were not forming any clear linkage between the filaments. The segmentation shown has a width of 73 nm x 101 nm and a height of 64 nm.

b-d Various examples of filament connections by putative ZapD proteins within the toroid. Same color code as in a. From left to right, the localization of the analyzed region, a close-up view of the structure of interest, different views of the crosslinkers and a schematic illustrating the interpretation of the data. The schematic (right) shows the localization of ZapD proteins (magenta) and FtsZ filaments (blue).

b. Lateral connection of two FtsZ filaments by a putative ZapD dimer. In this example, the attachment of each globular density or putative ZapD monomer was bound to each filament, allowing for lateral binding.

c. Putative ZapD connections stabilizing two filaments by a lateral interaction. Additional ZapD decorations attached to only one of the filaments appear to be available for other filament connections.

d. Multiple ZapD proteins can connect to filaments and stabilize the interaction. First, the two upper filaments are connected vertically by several putative ZapDs. The lower filament connects vertically in an oblique angle to the nearest neighboring filament. In the upper part, additional decorations or putative ZapD proteins would be available to establish further interactions forming a 3D mesh.

Earlier version: Cryo-ET imaging of ZapD-mediated FtsZ toroidal structures revealed a preferential vertical stacking and crosslinking of short ZapD filaments, which are also crosslinked laterally and diagonally, allowing for filament curvature and resulting in a toroidal structure observed for the first time following the interaction between FtsZ and one of its natural partners in vitro.

Revised version: Cryo-ET imaging of ZapD-mediated FtsZ toroidal structures revealed a preferential vertical stacking and crosslinking of short FtsZ filaments, which are also crosslinked laterally and diagonally, allowing for filament curvature and resulting in a toroidal structure observed for the first time following the interaction between FtsZ and one of its natural partners in vitro.

Earlier version: High concentrations of ZapD promote the structural reorganization of FtsZ polymers into straight bundles by tightly packed crosslinking from above

Revised version: High concentrations of ZapD promote the structural reorganization of FtsZ polymers into straight bundles by tightly packed crosslinking

Having characterized the higher order FtsZ structures promoted at equimolar ZapD binding densities, we addressed the effect of increasing the ZapD density on the structural organization of the ZapD-FtsZ polymers. At high concentrations of ZapD (typically 40 - 60 μM, representing a molar ratio of approximately one to four or six to FtsZ), we observed the formation of straight bundles with striated patterns between the FtsZ filaments (Supplementary Fig. 12), as well as the presence of some toroidal structures (Supplementary Fig. 6a). Here, the high concentration of ZapD molecules increased the number of links between the filaments and ultimately promoted the formation of straight bundles, indicating that the assembly of FtsZ-ZapD structures is a reversible process that strongly depends on the amount of ZapD proteins crosslinking the filaments. Toroids and curved bundles always coexist, but the predominance shifts from toroids to straight bundles at high ZapD concentrations. This shift suggests that the number of crosslinks between filaments can modulate the properties of the assembled higher order structures, resulting in the reorganization of the toroidal structures into straight bundles upon ZapD saturation (Supplementary Fig. 12). We then explored how higher densities of ZapD lead to the formation of straight bundles using cryo-ET.

Earlier version: Using the same approach to visualize the toroids, we collected tomograms of the straight bundles and extracted their isosurfaces (Fig. 5a, Supplementary Fig. 7b, Supplementary Fig. 13 and Supplementary Movie 3). The straight bundles are formed at high ZapD concentrations and consist of a highly organized stack of well-aligned FtsZ filaments (Fig. 5b). The connection of filaments by multiple putative ZapD proteins results in the straightening of the structure of the filaments10,50 (Fig. 5c). Interestingly, ZapD crosslink FtsZ filaments from the top, forming a row of ZapD proteins interacting with both filaments (Fig. 5d). Here too, structural analysis by subtomogram averaging was not possible due to the crowding, preferential orientation, and heterogeneity of the connections. The distance between ZapD proteins provided a mean value of 4.5 ± 0.5 nm (Supplementary Fig. 14a), consistent with the size of FtsZ proteins52. This observation suggests that most FtsZ proteins interact with the ZapD dimers that crosslink the filaments, in agreement with the ZapD enrichment in the bundles observed by sedimentation assays (Supplementary Fig. 5). In addition, we found that the presence of ZapD increased the distance between two FtsZ filaments connected by ZapD proteins compared with the spacing in the absence of ZapD, regardless of the amount of ZapD connections (from 5.9 ± 0.8 nm to 7.88 ± 2 nm in toroids and 7.6 ± 1.5 nm in straight bundles; Supplementary Fig.14b). This indicates that ZapD not only connects neighboring FtsZ filaments but spreads them apart, which could have important implications for the functionality of the Z ring.

Revised version: Using the same approach to visualize the toroids, we collected tomograms of the straight bundles and extracted their isosurfaces (Fig. 5a, Supplementary Fig. 7b, Supplementary Fig. 13 and Supplementary Movie 3). The straight bundles are formed at high ZapD concentrations and consist of a highly organized stack of well-aligned FtsZ filaments (Fig. 5b). The connection of filaments by multiple putative ZapD proteins results in the straightening of the structure of the filaments10,50 (Fig. 5c). Interestingly, ZapD crosslink FtsZ filaments vertically, forming a row of ZapD proteins interacting with both filaments (Fig. 5d). Here too, structural analysis by subtomogram averaging was not possible due to the crowding, preferential orientation, and heterogeneity of the connections. The distance between ZapD proteins provided a mean value of 4.5 ± 0.5 nm (Supplementary Fig. 14a), consistent with the size of FtsZ proteins52. This observation suggests that most FtsZ proteins interact with the ZapD dimers that crosslink the filaments, in agreement with the ZapD enrichment in the bundles observed by sedimentation assays (Supplementary Fig. 5). In addition, we found that the presence of ZapD increased the distance between two FtsZ filaments connected by ZapD proteins compared with the spacing in the absence of ZapD, regardless of the amount of ZapD connections (from 5.9 ± 0.8 nm to 7.88 ± 2 nm in toroids and 7.6 ± 1.5 nm in straight bundles; Supplementary Fig.14b). This indicates that ZapD not only connects neighboring FtsZ filaments but spreads them apart, which could have important implications for the functionality of the Z ring.

Formation of straight FtsZ bundles is driven by high ZapD crosslinking from above.

a. Representative tomographic slices of straight FtsZ bundles resulting from the interaction of FtsZ with high ZapD concentrations under working conditions (50 mM KCl, 50 mM Tris-Cl, 5 mM MgCl2 pH 7). The concentrations of FtsZ and ZapD were 10 µM and 60 µM, respectively, and 1 mM of GTP was added to trigger polymerization. The straight bundles were found only at high ZapD concentrations. The image is the average of five 0.86 nm thick tomographic slices (total thickness of 4.31 nm) of the reconstructed tomogram. Scale bars are 100 nm.

Earlier version: b. Isosurface of the straight bundles from the denoised tomographic volume. FtsZ filaments are colored in blue and putative ZapD connections in magenta. Three different views (top (left) and side views (middle, right)) are shown. Straight bundles are organized in a regular organization. Multiple bonds between filaments are formed from the top by putative ZapDs vertically crosslinking two FtsZ filaments with a regular spacing of 4.5 ± 0.5 nm between ZapD dimers. In addition, lateral connections were also found, connecting pairs of stabilized filaments to each other and eventually assembling the straight bundle.

Revised version: b. Isosurface of the straight bundles from the denoised tomographic volume. FtsZ filaments are colored in blue and putative ZapD connections in magenta. Three different views (top (left) and side views (middle, right)) are shown. Straight bundles are organized in a regular organization. Multiple bonds between filaments are formed vertically by putative ZapDs vertically crosslinking two FtsZ filaments with a regular spacing of 4.5 ± 0.5 nm between ZapD dimers. In addition, lateral connections were also found, connecting pairs of stabilized filaments to each other and eventually assembling the straight bundle.

c. Different views of one of the isolated filaments from the straight bundle. A side view of the filaments (middle) shows a spike-like structure regularly located at the top of the FtsZ filaments, connecting them vertically as observed in the top view (right).

d. Different close-up views of the filament structure shown in c. In the cross-section of the structure (middle, top), it is clearly visible that the ZapD proteins connect the two filaments vertically and from above, forming a bridge over them. A schematic of the proposed interaction (right, bottom) shows the position of putative ZapD dimers in this structure.

These results point out the relevant role of the number of connections between filaments (driven by ZapD binding) in the crosslinking mechanism (Fig. 6). The bundles formed at the high ZapD concentrations tested (up to four molar excess of ZapD to FtsZ) correspond to a binding saturation regime in which a ZapD dimer binds two FtsZ molecules connecting two filaments, straightening their structure and arranging them into large, highly organized straight bundles. In contrast, fewer connections provide the freedom to form toroidal structures. This observation confirms that a certain number of ZapD-driven bonds are optimal for maintaining the preferential curvature of FtsZ filaments and allowing toroid assembly, establishing a structural dependence of assembled bundles on the number of ZapD-mediated bonds.

The amount of ZapD connections modulates the spatial organization of FtsZ filaments into higher order structures.

Earlier version: (Top) Schematic of the higher order FtsZ structures formed in the presence of increasing concentrations of ZapD. (Bottom) Cryo-EM images of the structures shown in the schemes. In the absence of ZapD, FtsZ filaments can interact laterally to form double filaments upon GTP binding. At low concentrations of ZapD, only few ZapD-mediated bonds are formed, favoring the formation of small, curved bundles. At equimolar concentrations of ZapD and FtsZ, more ZapD-mediated bonds are formed, particularly from above the filaments, but also laterally and diagonally, allowing filament curvature and favoring the assembly of toroidal structures. At saturating ZapD concentrations, the filaments are straightened up by regular ZapD crosslinking from above, resulting in the formation of straight bundles. Overall, the assembly of higher order FtsZ structures depends on the number of vertical crosslinks through ZapD dimers. Some intermediate states are expected between the structures shown. Scale bars are 100 nm.

Revised version: (Top) Simplified scheme of the higher order FtsZ structures formed in the presence of increasing concentrations of ZapD. These schemes do not consider all the possible interactions between ZapD and FtsZ. (Bottom) Cryo-EM images of the structures shown in the schemes. In the absence of ZapD, FtsZ filaments can interact laterally to form double filaments upon GTP binding. At low concentrations of ZapD, only few ZapD-mediated bonds are formed, favoring the formation of small, curved bundles. At equimolar concentrations of ZapD and FtsZ, more ZapD-mediated bonds are formed, particularly from above the filaments, but also laterally and diagonally, allowing filament curvature and favoring the assembly of toroidal structures. At saturating ZapD concentrations, the filaments are straightened up by regular ZapD crosslinking vertically, resulting in the formation of straight bundles. Overall, the assembly of higher order FtsZ structures depends on the number of vertical crosslinks through ZapD dimers. Some intermediate states are expected between the structures shown. Scale bars are 100 nm.

Dimerization of ZapD is essential for the formation of organized higher-order FtsZ structures

Our results suggest that ZapD dimers connect two FtsZ filaments, although the role of dimerization has not yet been experimentally established41. We therefore investigated whether the dimerization of ZapD was essential for assembling straight bundles and toroidal structures. To this end, we produced a ZapD mutant (“mZapD”) by substituting three amino acids (R20, R116 and H140) involved in its dimerization by alanine, thus decreasing the stability of the dimerization interface. The dimerization of mZapD was tested by AUC analysis, which showed a percentage of monomeric protein that was not present in the wild-type ZapD (Supplementary Fig. 15a). mZapD can interact with unpolymerized FtsZ-GDP as evidenced by fluorescence anisotropy and FCS (Supplementary Fig. 15b and c). mZapD is still able to promote FtsZ bundle formation, showing a ∼50% lower turbidity signal than for wild-type ZapD even at higher concentrations (Supplementary Fig. 15d). We also visualized these bundles by cryo-EM and observed the formation of thinner bundles (Supplementary Fig. 16b and c). However, toroidal structures and straight bundles were absent in these samples even at high concentrations of mZapD. This indicates that a stable dimerization interface is required to assemble complex ZapD-mediated structures (Supplementary Fig. 16a).

These results demonstrated that dimerization of ZapD is essential to promote the assembly of organized high-order FtsZ structures such as toroids and straight bundles, highlighting the structural importance of stable ZapD-driven connections between filaments.

Discussion

In this study, we used biochemical reconstitution, cryo-EM and cryo-ET to gain new insights into the structure and assembly of the bacterial divisome. Filament crosslinking by Zap proteins plays a crucial role in the assembly and stabilization of the Z ring. However, the underlying mechanisms of this process remain to be elucidated. To better understand these mechanisms, we investigated the interaction between FtsZ and ZapD, one of the stabilizers of the division ring.

Earlier version: Our results indicate that ZapD crosslinks FtsZ in solution to form 3D toroidal structures characterized by short, discontinuous FtsZ filaments that are vertically stacked and crosslinked, as well as diagonally and laterally crosslinked, as revealed by cryo-ET. The relative concentrations of FtsZ and ZapD dictate the assembly geometry. Remarkably, FtsZ toroids emerge as the dominant structure when both proteins are present in equimolar ratios, whereas straight bundles are assembled at high concentrations of ZapD. Our study highlights the importance of the balance between vertical and lateral or diagonal crosslinking to allow for the filament curvature observed in the toroids, which is absent in the straight bundles where crosslinking from the top is dominant.

Revised version: Our results indicate that ZapD crosslinks FtsZ in solution to form 3D toroidal structures characterized by short, discontinuous FtsZ filaments that are vertically stacked and crosslinked, as well as diagonally and laterally crosslinked, as revealed by cryo-ET. The relative concentrations of FtsZ and ZapD dictate the assembly geometry. Remarkably, FtsZ toroids emerge as the dominant structure when both proteins are present in equimolar ratios, whereas straight bundles are assembled at high concentrations of ZapD. Our study highlights the importance of the balance between vertical and lateral or diagonal crosslinking to allow for the filament curvature observed in the toroids, which is absent in the straight bundles where vertical crosslinking is dominant.

The structural characterization of the toroids by cryo-ET allowed us to resolve the overall structure and identify putative ZapD dimers that crosslink the filaments. However, we were unable to obtain detailed structural information about the ZapD connectors due to the heterogeneity and density of the toroidal structures, which showed significant variability in the conformations of the connections between the filaments in all directions. These results are consistent with the observation that ZapD interacts with FtsZ through its central hub, which provides additional spatial freedom to connect other filaments in different conformations. This flexibility allows different filament organizations and contributes to structural heterogeneity. In addition, these results suggest that these crosslinkers can act as modulators of the dynamics of the ring structure, spacing filaments apart and allowing them to slide in an organized manner33,37,53,54. The ability of FtsZ to treadmill directionally34,55–58, together with the parallel or antiparallel arrangement of short, transiently crosslinked filaments, is considered essential for the functionality of the Z ring and its ability to exert constrictive force34,36–38,50. Thus, Zap proteins can play a critical role in ensuring correct filament placement and stabilization, which is consistent with the toroidal structure formed by ZapD.

Our observation of toroidal FtsZ structures promoted by ZapD in solution is consistent with the observation of pre-curved FtsZ protofilaments in circular assemblies attached to model membranes in various in vitro systems55,56,59. Additionally, FtsZ can form ring-shaped structures with diameters ranging from 100 to 200 nm on surfaces, becoming more pronounced when adsorbed to lipid, carbon, or mica surfaces52,60,61, or in the presence of molecular crowders like methylcellulose14. Cryo-EM of concentrated (50 μM) FtsZ from Bacillus subtilis with GMPCPP revealed that these protofilaments often coalesce into spirals or toroids, forming large aggregates31. These toroidal structures correspond to the curved conformations of FtsZ polymers observed in various bacterial species, which are thought to contribute to the generation of constriction forces50. In Escherichia coli, approximately 5,000 FtsZ monomers (about one third of the intracellular concentration of FtsZ) could circle the cell 6-8 times62, suggesting a discontinuous toroidal assembly. ZapD, present at approximately 500 molecules per cell19, represents only a small fraction of the FtsZ molecules and favors the formation of toroidal structures in the absence of interactions with other FtsZ-associated proteins. Higher concentrations of ZapD would be required to form straight bundles.

The persistence length and curvature of FtsZ filaments are optimized to form bacterial-sized ring structures7,50. ZapD helps to stabilize these toroids by crosslinking and increasing the spacing between FtsZ filaments from 5.9 ± 0.8 nm to 7.9 ± 2 nm in toroids and 7.6 ± 1.5 nm in straight bundles (Supplementary Fig. 14b). This increased spacing promotes a more dynamic organization, providing functional flexibility in bacteria. The linker connecting ZapD to FtsZ can also modulate filament spacing and curvature31,32,63. We observe a mixture of curvatures in the internal organization of the toroids. While FtsZ filaments have a preferred curvature, the density of ZapD connections forces the structure to adapt and align with neighboring filaments, thereby affecting FtsZ curvature. However, the precise molecular mechanism linking ZapD binding and polymer curvature remains to be elucidated.

Earlier version: One of the most significant findings of this study is that the amount of ZapD bound to FtsZ filaments greatly influences the structural organization of the resulting higher order polymer. At lower concentrations of ZapD, which are necessary for bundling (equimolar mixtures of ZapD and FtsZ), toroids are the most prominent structures, containing one ZapD dimer for every four to six FtsZ molecules. When the concentration of ZapD increases, the density of ZapD binding in the polymers rises to one ZapD dimer for every two FtsZ molecules. This leads to a reorganization of the polymer bundles, resulting in straight structures that predominantly crosslink the FtsZ filaments, a configuration not seen at lower ZapD concentrations (Fig. 5). The increase in the number of ZapD-FtsZ contacts likely reduces the flexibility needed to form toroidal structures, compelling the filaments to adopt a larger, straight bundle conformation. We observed slight differences in the spacing between filaments (Supplementary Fig. 14b), with a broader range of distances between filaments in the toroids compared to the straight bundles. This makes sense, as ZapD is the molecule linking the filaments in both structures, although the amount of bound ZapD differs, allowing for greater variability in spacing within regions with a relatively fewer ZapD-FtsZ connections.

Revised version: One of the most significant findings of this study is that the amount of ZapD bound to FtsZ filaments greatly influences the structural organization of the resulting higher order polymer. At lower concentrations of ZapD, which are necessary for bundling (equimolar mixtures of ZapD and FtsZ), toroids are the most prominent structures, with a binding stoichiometry of ZapD in the polymer of 0.3-0.4 moles of ZapD per mole of FtsZ (equivalent to one ZapD dimer for every four to six FtsZ molecules). As the concentration of ZapD increases, the binding stoichiometry of ZapD in the polymers increases to 1.1 moles of ZapD per mole of FtsZ (one ZapD dimer for every two FtsZ molecules). This leads to a reorganization of the polymer bundles, resulting in straight structures that predominantly crosslink the FtsZ filaments, a configuration not seen at lower ZapD concentrations (Fig. 5). The increase in the number of ZapD-FtsZ contacts likely reduces the flexibility needed to form toroidal structures, compelling the filaments to adopt a larger, straight bundle conformation. We observed slight differences in the spacing between filaments (Supplementary Fig. 14b), with a broader range of distances between filaments in the toroids compared to the straight bundles. Thus, ZapD contributes to the variability of spacing based on the number of connections formed.

We therefore hypothesize that the assembly of functional, curved FtsZ macrostructures occurs only within a specific stoichiometric range of ZapD-FtsZ interactions. An increase in crosslinking at higher ZapD concentrations appears to cause the FtsZ polymer to form rigid, straight bundles (Fig. 6). Additionally, only a robust ZapD dimer can form FtsZ toroids and straight filaments, indicating that a certain binding strength is necessary to bundle the filaments and maintain their structure. These observations, along with the high conformational variability found in ZapD connections, suggest that ZapD can modulate the behavior of the entire structure through a concentration-dependent mechanism. This structural reorganization highlights the polymorphic nature of the assemblies formed by FtsZ38,64.

We also demonstrated that ZapD-mediated FtsZ toroids replicate some of the structural features of the Z ring architecture observed in vivo. Based on super-resolution imaging and cryo-electron tomography (cryo-ET), the current model of the Z ring describes a somewhat discontinuous and heterogeneous structure composed of randomly overlapping filaments33,45,65. These filaments arrange into a belt-like macromolecular entity, containing nodes with a higher density of dispersed filaments confined to a toroidal zone approximately 80–100 nm wide (toroidal height) and about 40–60 nm thick, located 13-16 nm below the inner membrane33,39,45,47,65–67. The toroidal structures promoted by ZapD exhibit dimensions within this range. However, their thickness and multilayering are significantly greater, likely due to the absence of other cellular components competing for binding with FtsZ. The Z ring model also suggests that FtsZ filaments are weakly associated with one another through protein factors such as Zaps and weak FtsZ-FtsZ interactions, which may be crucial for their functionality7,31,32,68. Our findings indicate that one of the Zap proteins can stabilize and enhance associations between neighboring FtsZ filaments. Studying macrostructures stabilized by natural crosslinkers is vital for understanding the molecular function of the division machinery.

Earlier version: We believe that the intrinsic features of the toroidal structures share commonalities with the bacterial division ring. Despite the differences and limitations that may arise from an in vitro approach, the structures observed following ZapD crosslinking of FtsZ filaments may reveal inherent features that occur in vivo. The current model of the division ring consists of an array of filaments loosely connected by crosslinkers at the center of the cell, forming a ring. This model is consistent with our findings, although many questions remain regarding the structural organization of the Z ring within the cell. ZapD binds to FtsZ from above, allowing either ZapD or FtsZ to interact with the plasma membrane. In straight bundles, this facilitates the stacking of straight FtsZ filaments, while for toroids, ZapD can also bind FtsZ filaments diagonally. This less compact arrangement could allow bending of the FtsZ filaments and adjustment of the toroid size.

Revised version: We believe that the intrinsic features of the toroidal structures share commonalities with the bacterial division ring. Despite the differences and limitations that may arise from an in vitro approach, the structures observed following ZapD crosslinking of FtsZ filaments may reveal inherent features that occur in vivo. The current model of the division ring consists of an array of filaments loosely connected by crosslinkers at the center of the cell, forming a ring. This model is consistent with our findings, although many questions remain regarding the structural organization of the Z ring within the cell. ZapD binds to FtsZ vertically, allowing either ZapD or FtsZ to interact with the plasma membrane. In straight bundles, this facilitates the stacking of straight FtsZ filaments, while for toroids, ZapD can also bind FtsZ filaments diagonally. This less compact arrangement could allow bending of the FtsZ filaments and adjustment of the toroid size.

Earlier version: The binding of ZapA, a key member of the FtsZ-associated proteins family, to FtsZ in equimolar mixtures induces the effective alignment and straightening of FtsZ filaments tethered to lipid bilayers by FtsA protein, as demonstrated through high-resolution fluorescence microscopy and precise quantitative image analysis69. This interaction reveals a structure comprising one ZapA tetramer for every four FtsZ molecules in the polymer, highlighting the organized nature of these interactions. It is noteworthy that these straight, parallel polymers—manifested when ZapA occupies all binding sites—bear striking resemblance to the straight bundles identified in this study at high concentrations of ZapD, indicative of binding saturation. Understanding how the associative states of ZapA (as tetramers) and ZapD (as dimers), together with membrane tethering, influence the predominant structures formed in both systems is essential. The complexity of the division system raises important questions about the interaction dynamics between FtsZ and the plasma membrane. The competitive nature of the division components to engage with FtsZ and modulate its functionality remains to be thoroughly elucidated. It is important to note that FtsA and ZipA have a greater affinity for the C-terminal domain of FtsZ than ZapD. Our cryo-ET data on straight bundles provide new perspectives on how ZapD-FtsZ structures can effectively bind to the plasma membrane; in particular, the C-terminal domains of parallel FtsZ filaments are oriented upward, allowing direct membrane binding or interaction with ZapDs that reinforce these filaments from above, rather than from the side, as previously suggested41.

Revised version: The binding of ZapA, a key member of the FtsZ-associated proteins family, to FtsZ in equimolar mixtures induces the effective alignment and straightening of FtsZ filaments tethered to lipid bilayers by FtsA protein, as demonstrated through high-resolution fluorescence microscopy and precise quantitative image analysis69. This interaction reveals a structure comprising one ZapA tetramer for every four FtsZ molecules in the polymer, highlighting the organized nature of these interactions. It is noteworthy that these straight, parallel polymers—manifested when ZapA occupies all binding sites—bear striking resemblance to the straight bundles identified in this study at high concentrations of ZapD, indicative of binding saturation. Understanding how the associative states of ZapA (as tetramers) and ZapD (as dimers), together with membrane tethering, influence the predominant structures formed in both systems is essential. The complexity of the division system raises important questions about the interaction dynamics between FtsZ and the plasma membrane. The competitive nature of the division components to engage with FtsZ and modulate its functionality remains to be thoroughly elucidated. It is important to note that FtsA and ZipA have a greater affinity for the CCTP of FtsZ than ZapD. Our cryo-ET data on straight bundles provide new perspectives on how ZapD- FtsZ structures can effectively bind to the plasma membrane; in particular, the CCTP of parallel FtsZ filaments are oriented upward, allowing direct membrane binding or interaction with ZapDs that reinforce these filaments from above, rather than from the side, as previously suggested41.

In conclusion, this study extends the understanding of the intricate structural organization and filament crosslinking of FtsZ polymers facilitated by ZapD, which is critical for the formation and maintenance of the division ring. The revelation that ring morphology depends on the binding stoichiometry between FtsZ and the crosslinking ZapD is invaluable, particularly as we integrate this knowledge into current reconstitution assays in membrane systems. This knowledge is central to unraveling the role of FtsZ-associated proteins in coordinating and maintaining division ring formation. In addition, our findings could greatly enhance the understanding of how polymeric cytoskeletal networks are remodelled during essential cellular processes such as cell motility and morphogenesis. Although conventional wisdom points to molecular motors as the primary drivers of filament remodeling through energy consumption, there is increasing evidence that there are alternative mechanisms that do not rely on such energy, instead harnessing entropic forces via diffusible crosslinkers70. This approach may also be applicable to ZapD and FtsZ polymers, suggesting a promising avenue for optimizing conditions in the reverse engineering of the division ring to enhance force generation in minimally reconstituted systems aimed at achieving autonomous cell division.

Methods

Protein purification

Removed in the revised version: ZapD protein has been overproduced and purified following the procedure previously described in19 with some modifications. The bacterial strain was an E. coli BL21 (DE3) transformed with pET28b-h10-smt3-ZapD fusion plasmid19. The cells were sonicated for 6-8 cycles of 20 secs and centrifuged at 40.000 rpm in MLA 80 rotor (Beckman coulter) for 45 min. The protein was eluted by chromatography in 5-ml HisTrap (GE Healthcare) and digested with Smt3-specific protease Ulp1 (SUMO protease encoded in Rosetta pTB145), at 1:1 molar relation. Digestion proceeded for two dialysis of 1 hour at 4°C to avoid protein precipitation. ZapD and His-Smt3 were separated by chromatography in 5-ml HisTrap (GE Healthcare), then the protein was eluted in a 5-ml HiTrap Desalting column (Healtcare) to eliminate completely possible traces of free phosphates in the buffer. Final concentration of ZapD was determined by absorbance at 280 nm using an extinction coefficient (ε) of 25230 M-

Earlier version: 1 cm-1 and ZapD purity was checked by SDS-polyacrylamide gel electrophoresis. Protein fractions were frozen in buffer 50 mM KCL, 50 mM Tris-CL, 10 mM MgCl2, 0.2 mM DTT and 2% glycerol. ZapD mutant (mZapD) was purified following the same procedure using the corresponding plasmid. E. coli FtsZ was overproduced and purified following the calcium precipitation method described previously48. Briefly, E. coli BL21 (DE3) pLysS cells were transformed with pET11b-FtsZ, grown in LB medium and selected with Ampicillin 100 μg/mL. After induction and growth, the pellet was resuspended in PEM buffer (50 mM PIPES-NaOH, pH 6.5, 5 mM MgCl2, 1 mM EDTA) and disrupted using a tip sonicator for 3-4 cycles. The lysate was then separated by centrifugation for 30 min at 20,000 x g at 4 °C, and the supernatant was mixed with 1 mM GTP, 20 mM CaCl2 and incubated at 30 °C for 15 min to induce FtsZ polymerization and bundling. Subsequently, the FtsZ bundles were pelleted by centrifugation for 15 min at 20,000 x g at 4 °C, and the pellet was resuspended in PEM buffer and centrifuged again for 15 min at 20,000 x g, 4 °C, collecting the supernatant. Precipitation and resuspension steps were repeated to improve the purification. The buffer was then exchanged using an Amicon Ultra-0.5 centrifugal filter unit 50 kDa (Merck KGaA). FtsZ purity was checked by SDS-polyacrylamide gel electrophoresis and concentration was determined by absorbance at 280 nm using an extinction coefficient (ε) of 14000 M-1 cm-1 48. Protein solutions were aliquoted, frozen in liquid nitrogen, and stored at –80 °C until further use.

Revised version: ZapD protein has been overproduced and purified following the procedure previously described in19 with some modifications. The bacterial strain was an E. coli BL21 (DE3) transformed with pET28b-h10-smt3-ZapD fusion plasmid19. The cells were sonicated for 6-8 cycles of 20 secs and centrifuged at 40.000 rpm in MLA 80 rotor (Beckman coulter) for 45 min. The protein was eluted by chromatography in 5-ml HisTrap (GE Healthcare) and digested with Smt3-specific protease Ulp1 (SUMO protease encoded in Rosetta pTB145), at 1:1 molar relation. Digestion proceeded for two dialysis of 1 hour at 4°C to avoid protein precipitation. ZapD and His-Smt3 were separated by chromatography in 5-ml HisTrap (GE Healthcare), then the protein was eluted in a 5-ml HiTrap Desalting column (Healthcare) to eliminate completely possible traces of free phosphates in the buffer. Final concentration of ZapD was determined by absorbance at 280 nm using an extinction coefficient (ε) of 25230 M-1 cm-1 and ZapD purity was checked by SDS-polyacrylamide gel electrophoresis. Protein fractions were frozen in buffer 50 mM KCl, 50 mM Tris-CL, 10 mM MgCl2, 0.2 mM DTT and 2% glycerol. ZapD mutant (mZapD) was purified following the same procedure using the corresponding plasmid. E. coli FtsZ was overproduced and purified following the calcium precipitation method described previously48. Briefly, E. coli BL21 (DE3) pLysS cells were transformed with pET11b-FtsZ, grown in LB medium and selected with Ampicillin 100 μg/mL. After induction and growth, the pellet was resuspended in PEM buffer (50 mM PIPES-NaOH, pH 6.5, 5 mM MgCl2, 1 mM EDTA) and disrupted using a tip sonicator for 3-4 cycles. The lysate was then separated by centrifugation for 30 min at 20,000 x g at 4 °C, and the supernatant was mixed with 1 mM GTP, 20 mM CaCl2 and incubated at 30 °C for 15 min to induce FtsZ polymerization and bundling. Subsequently, the FtsZ bundles were pelleted by centrifugation for 15 min at 20,000 x g at 4 °C, and the pellet was resuspended in PEM buffer and centrifuged again for 15 min at 20,000 x g, 4 °C, collecting the supernatant. Precipitation and resuspension steps were repeated to improve the purification. The buffer was then exchanged using an Amicon Ultra-0.5 centrifugal filter unit 50 kDa (Merck KGaA). FtsZ purity was checked by SDS- polyacrylamide gel electrophoresis and concentration was determined by absorbance at 280 nm using an extinction coefficient (ε) of 14000 M-1 cm-1 48. Protein solutions were aliquoted, frozen in liquid nitrogen, and stored at -80 °C until further use.

Plasmid Design and Molecular Cloning

To generate the ZapD mutant (mZapD), seamless cloning method was used according to the provider’s protocol (ThermoFisher Scientific /Invitrogen GeneArt™ Seamless Cloning and Assembly Enzyme Mix (A14606)) using the plasmid pET28b-h10-smt3-ZapD and the primers listed in (Supplementary Table 1). All enzymes for cloning were from Thermo Fisher Scientific (Waltham, MA, USA). Briefly, DNA fragments were amplified with Phusion High-Fidelity DNA Polymerase andorigo primers (Sigma–Aldrich, St. Louis, MO, USA). Then, PCR products were treated with DpnI and combined using GeneArt Seamless Cloning and Assembly Enzyme Mix. The plasmid was propagated in E. coli OneShot TOP10 (Thermo Fisher Scientific) and purified using NucleoBond Xtra Midi kit (Macherey-Nagel GmbH, Duren, Germany). Directed site mutagenesis was made by substituting three amino acids in the ZapD sequence by Alanine (R20, R116, H140). The plasmid was then verified using Sanger Sequencing Service (Microsynth AG, Balgach, Switzerland).

Protein labelling

Covalent labelling of FtsZ with Alexa 488 the amino groups of N-terminal amino acid residue with Alexa Fluor 488 carboxylic acid succinimidyl ester dye following the procedure previously described13. ZapD and mZapD were labelled with ATTO-647N carboxylic acid succinimidyl ester dye in the amino group. Before the reaction, ZapD was dialyzed in 20 mM HEPES, 50 mM KCl, 5 mM MgCl2, pH 7.5 and the probe was dissolved in Dimethylsulfoxide (DMSO). The reaction was allowed to proceed for 35-60 min at RT and stopped with 10 % Tris-HCl 1 M. The free dye was separated from labelled protein by a Hi- TRAP Desalting column (GE Healthcare). The final degree of labelling of FtsZ and ZapD was calculated from the molar absorption coefficients of the protein and the dye. It was around 0.5 moles of probe per mole of FtsZ and around 0.3/0.4 moles of dye per mole of ZapD.

Analytical ultracentrifugation (AUC); sedimentation velocity (SV) and sedimentation equilibrium (SE)

Earlier version: Sedimentation velocity assays were performed to detect the homogeneity and association state of individual proteins and the stoichiometry of the formed protein-protein complexes. In brief, the experiments were carried out at 43-48 Krpm in an Optima XL-I analytical ultracentrifuge, equipped with UV–VIS absorbance and Raleigh interference detection systems. The sedimentation coefficient distributions were calculated by least-squares boundary modeling of sedimentation velocity data using the c(s) method71 as implemented in the SEDFIT program. The s-values of the present species were corrected to standard conditions (pure water at 20°C, and extrapolated to infinite dilution) to obtain the corresponding standard s-values (s20,w) using the program SEDNTERP72. Multi-signal sedimentation velocity (MSSV) data were globally analyzed by SEDPHAT software73 using the “multi-wavelength discrete/continuous distribution analysis” model, to determine the spectral and diffusion deconvoluted sedimentation coefficient distributions, ck(s), from which the stoichiometry of protein complexes can be derived74. Sedimentation equilibrium of ZapD was carried out to confirm the association state of the protein in the same experimental conditions and concentration range tested by sedimentation velocity (2-30 µM). Short columns (100 µL) SE experiments were carried out at 14000 and 10000 rpm. Weight-average buoyant molecular weights were obtained by fitting a single-species model to the experimental data using the HeteroAnalysis program75, once corrected for temperature and buffer composition with the program SEDNTERP72.

Revised version: Sedimentation velocity assays were performed to detect the homogeneity and association state of individual proteins and the stoichiometry of the formed protein-protein complexes. In brief, the experiments were carried out at 43-48K rpm in an Optima XL-I analytical ultracentrifuge, equipped with UV–VIS absorbance and Raleigh interference detection systems. The sedimentation coefficient distributions were calculated by least- squares boundary modeling of sedimentation velocity data using the c(s) method71 as implemented in the SEDFIT program. The s-values of the present species were corrected to standard conditions (pure water at 20°C, and extrapolated to infinite dilution) to obtain the corresponding standard s-values (s20,w) using the program SEDNTERP72. Multi-signal sedimentation velocity (MSSV) data were globally analyzed by SEDPHAT software73 using the “multi-wavelength discrete/continuous distribution analysis” model, to determine the spectral and diffusion deconvoluted sedimentation coefficient distributions, ck(s), from which the stoichiometry of protein complexes can be derived74. Sedimentation equilibrium of ZapD was carried out to confirm the association state of the protein in the same experimental conditions and concentration range tested by sedimentation velocity (2-30 µM). Short columns (100 µL) SE experiments were carried out at 14000 and 10000 rpm. Weight-average buoyant molecular weights were obtained by fitting a single-species model to the experimental data using the HeteroAnalysis program75, once corrected for temperature and buffer composition with the program SEDNTERP72.

Turbidity assay

Turbidity of protein samples was collected by measuring the absorbance at 350 nm in a TECAN plate reader (Tecan Group Ltd., Mannedorf, Switzerland). All samples reached a final volume of 50 μL in a 364-Well Flat-Bottom Microplate (UV-Star, Greiner Bio- One GmbH) before measuring the absorbance. Different concentrations of FtsZ and ZapD were premixed in the well-plate and measured before addition of GTP to extract subsequently the individual blank values. Concentrations of FtsZ and ZapD varied from 0 to 80 μM and buffer conditions are specified in each graph and the caption of the figures ranging from 50 to 500 mM KCl, 6 to 8 pH, always supplemented with 50 mM Tris-Cl and 5 mM MgCl2. Manual mixing was performed after addition of GTP and orbital oscillations for 5 sec in the TECAN were made prior to any measurement to avoid sedimentation of the samples. FtsZ and ZapD alone do not generate significant differences with the blank. Time measurements were taken as specified in each condition. Reported values are the average of 3-12 independent measurements ± Standard deviation.

GTPase activity of FtsZ

GTPase activity of FtsZ was measured by quantification of the inorganic phosphate with a colorimetric assay (BIOMOL GREENÒ kit from ENZO life sciences) for two minutes. 5 μM FtsZ was used in our standard buffer (5 mM MgCl2, 50 mM Tris-HCl, 50 mM KCl, pH 7) or buffers at higher KCl concentrations (50-500 mM KCl) and polymerization was triggered by 1 mM GTP. ZapD was added at different concentrations and premixed with FtsZ before addition of GTP. 13 μL fractions were added to a 96-Well Flat-Bottom Microplate (UV-Star, Greiner Bio-One GmbH) every 20 sec after addition of GTP and mixed with the Reaction buffer reaching 50 μL and 100 μL of BIOMOL GREEN reagent, to stop the reaction. After stopping the reaction, samples were incubated for 10 min at RT and the absorbance was measured at 620 nm in a Varioskan Flash plate reader (Thermo Fisher Scientific, MA, USA). Concentrations of inorganic Phosphate were calculated from a phosphate standard curve, while the GTPase activity reaction rate (V, mol P/mol FtsZ/ min) was determined from the slope of the linear part of phosphate accumulation curves.

FtsZ sedimentation assay

Earlier version: Purified ZapD (1-30 μM) was added to purified FtsZ (5 μM) in the working buffer (50 mM KCl, 50 mM Tris-HCl, 5 mM MgCl2), and GDP or GTP (1 mM) was added last to trigger FtsZ polymerization. The reaction mixtures with a final volume of 100 μl were processed at room temperature and centrifuged at low speed (10.000 rcf) using a table top centrifuge. At that point, 90 μl of supernatant was carefully collected and loaded in clean tubes with 1x loading dye. The rest of the supernatant was discarded and the pellets were resuspended in the original reaction volume buffer plus 1× loading dye (final concentration). The supernatants (10 μl) and pellets (10 μl) were resolved in a SDS-PAGE gel. The amount of ZapD and FtsZ were estimated by image analysis using ImageJ.

Revised version: Purified ZapD (1-30 μM) was added to purified FtsZ (5 μM) in the working buffer (50 mM KCl, 50 mM Tris-HCl, 5 mM MgCl2), and GDP or GTP (1 mM) was added last to trigger FtsZ polymerization. The reaction mixtures with a final volume of 100 μl were processed at room temperature and centrifuged at low speed (10.000 rcf) using a table top centrifuge. At that point, 90 μl of supernatant was carefully collected and loaded in clean tubes with 1x loading dye. The rest of the supernatant was discarded and the pellets were resuspended in the original reaction volume buffer plus 1× loading dye (final concentration). The supernatants (10 μl) and pellets (10 μl) were resolved in a SDS-PAGE gel. The amount of FtsZ was estimated by image analysis using ImageJ.

Preparation of EM grids

Cryo-EM grids were plunge frozen with a Vitrobot Mk.IV (Thermo Fischer Scientific) using 3 µL of the samples applied to previously glow discharged R 2/1 Holey Carbon Cu 200 mesh EM grids (Quantifoil). Samples were 10 µM FtsZ with or without ZapD or mZapD at different concentrations specified in each case (0 – 60 µM). Proteins were mixed in our working buffer containing 50 mM, 5 mM MgCl2, 50 mM Tris-HCl, pH 7. Samples were incubated for 2 minutes after the addition of 1 mM GTP to trigger polymerization. The Vitrobot was set to 4° C, 100% humidity, blot time 3 s, blot force 3. Whatman no. 1 filter paper was used for blotting and liquid ethane kept at liquid nitrogen temperatures was used as a cryogen for vitrification.

Cryo-EM and cryo-ET

Cryo-EM/ET data were acquired on two microscopes as follows. Cryo- EM micrographs were acquired within SerialEM76 on a Talos Arctica transmission electron microscope (Thermo Fisher Scientific) operated at 200 kV, equipped with a Falcon III (Thermo Fisher Scientific) direct electron detector operated in integrating mode. Images were recorded at 73,000x magnification (pixel size 1.997 Å) and 92,000x magnification (pixel size 1.612 Å) at -2.5 µm to -5 µm target defocus with an approximate total electron dose of 60 electrons / Å². Cryo-EM micrographs and cryo-ET tilt series were acquired with SerialEM on a Titan Krios G2 transmission electron microscope (Thermo Fisher Scientific) operated at 300 kV, equipped with a FEG, post-column energy filter (Gatan) and a K3 camera (Gatan) operated in electron counting mode. Micrographs were recorded at 42,000x magnification (pixel size 2.154 Å) at - 5 µm target defocus with an approximate total electron dose of 60 electrons / Å². Tilt series were recorded at 42,000x magnification (pixel size 2.154 Å) at -5 µm target defocus with an approximate total electron dose of 100-140 electrons / Å². Acquisition was performed using a dose-symmetric tilt scheme, a tilt range of +/-60°, an angular increment of 2°.

Tomogram reconstruction

Tilt series preprocessing was performed using the TOMOMAN package (https://github.com/williamnwan/TOMOMAN). In brief, MotionCor277 was used to align individual frames, followed by cumulative dose-weighting using an exposure-dependent attenuation function, as described in78. Dose-weighted tilt series were aligned using IMOD79 either by employing gold fiducials (if available) or by patch-tracking, and binned tomograms (pixel size 8.616 Å) were generated using weighted back-projection. Stacks were split into odd and even frames to generate half-set tomograms which were used for training and subsequent denoising in cryoCARE80.

Earlier version: Denoised tomograms were used directly for segmentation of the toroids, straight bundles and individual filaments by cropping an area of interest and displaying the volume as an isosurface in USCF ChimeraX81. Thresholds for volume extraction were 80-100 and the color used were #00AEEF (light blue) for FtsZ and #EC008C (magenta) for ZapD isosurfaces. Small particles were removed by using the “dust” function with a size of 10. Different perspectives and zooms of the structures were manually adjusted for each figure in ChimeraX. To highlight connections between filaments, the corresponding parts with putative ZapD connections were manually cropped from the volume using the Volume Eraser tool in USCF Chimera82. Both cropped and original isosurfaces were superimposed to show the colored putative ZapDs in the structure. The missing wedge effect induces an elongation by a factor of 2 along the Z-axis. This elongation is observed in the filaments of the 1:0 ratio toroids, whereas the elongation observed in the filaments of the 1:1 and 1:6 ratio toroids exceeds the missing wedge.

Revised version: Denoised tomograms were used directly for segmentation of the toroids, straight bundles and individual filaments by cropping an area of interest and displaying the volume as an isosurface in USCF ChimeraX81. Thresholds for volume extraction were 80-100 and the color used were #00AEEF (light blue) for FtsZ and #EC008C (magenta) for ZapD isosurfaces. Small particles were removed by using the “dust” function with a size of 10. Different perspectives and zooms of the structures were manually adjusted for each figure in ChimeraX. To highlight connections between filaments, the corresponding parts with putative ZapD connections were manually cropped from the volume using the Volume Eraser tool in USCF Chimera82. Both cropped and original isosurfaces were superimposed to show the colored putative ZapDs in the structure. The missing wedge effect induces an elongation by a factor of 2 along the Z-axis. This elongation is observed in the filaments of the 1:0 ratio toroids, whereas the elongation observed in the filaments of the 1:1 and 1:6 ratio toroids exceed the missing wedge.

All micrographs and slices through tomograms were visualized using IMOD. Isosurface renderings of toroids, straight bundles and individual filaments were displayed using USCF ChimeraX and USCF Chimera.

Fluorescence anisotropy

Anisotropy measurements were performed using a TECAN plate reader (Tecan Group Ltd., Mannedorf, Switzerland). Excitation and emission wavelengths were 625 nm and 680 nm, respectively. ZapD or mZapD labelled with ATTO 647N were used as fluorescence tracer with a final concentration of 150 nM of ATTO-647N and supplemented with unlabeled ZapD reaching a concentration of 5 μM. FtsZ was added at increasing concentrations to analyze their interaction. Binding affinities (apparent KD) were determined by fitting the Hill equation to the normalized anisotropy data. Each condition was measured in three independent samples.

Fluorescence correlation spectroscopy (FCS)

Earlier version: FCS measurements were performed using a PicoQuant MicroTime200 system equipped with an Olympus 60x, NA1.2 UPlanApo water immersion objective. A pulsed 636 nm diode laser was used to excite fluorescence of Atto647N-labelled ZapD, mZapD or free Atto647N carboxylate (for calibration). Three measurements of 60 s each were performed per sample at room temperature (21°C), and three samples per condition were measured. The repetition rate of the pulsed laser was 26.7 MHz, and the average power was 1 µW at the objective back pupil. Fluorescence was collected through the same objective, spatially filtered through a 50 µm diameter pinhole, and spectrally filtered through a 690/70 nm bandpass filter before being split by a neutral beam splitter onto two avalanche photodiodes (Excelitas SPCM-AQRH-14-TR). Photon count events were digitized using a PicoQuant TimeHarp 260 Nano TCSPC card. Time-correlated single photon counting information was used for monitoring data quality during measurements, but not in further analysis. As especially in samples with FtsZ ZapD, and GTP present simultaneously we occasionally saw large, bright particles that are difficult to treat in FCS analysis, data was subjected to a burst removal algorithm similar to Margineanu et al. 83 and only “non-burst” data was used in correlation analysis to obtain statistics representative of the large number of small oligomer particles, ignoring rare large ones. Cross-correlation functions of the two detectors were calculated and fitted with a standard model for 3D diffusion in with rapid dye blinking: With amplitude G0, diffusion time τd, point spread function aspect ratio S, and blinking parameters FB” and τB”. Custom software was written in Python for burst removal, correlation, and fitting, based on tttrlib 0.0.19 (https://github.com/Fluorescence-Tools/tttrlib). The software is in ongoing development, and available upon request.

Revised version: FCS measurements were performed using a PicoQuant MicroTime200 system equipped with an Olympus 60x, NA1.2 UPlanApo water immersion objective. A pulsed 636 nm diode laser was used to excite fluorescence of Atto647N-labelled ZapD, mZapD or free Atto647N carboxylate (for calibration). Three measurements of 60 s each were performed per sample at room temperature (21°C), and three samples per condition were measured. The repetition rate of the pulsed laser was 26.7 MHz, and the average power was 1 µW at the objective back pupil. Fluorescence was collected through the same objective, spatially filtered through a 50 µm diameter pinhole, and spectrally filtered through a 690/70 nm bandpass filter before being split by a neutral beam splitter onto two avalanche photodiodes (Excelitas SPCM-AQRH-14-TR). Photon count events were digitized using a PicoQuant TimeHarp 260 Nano TCSPC card. Time-correlated single photon counting information was used for monitoring data quality during measurements, but not in further analysis. Data was subjected to a burst removal algorithm similar to Margineanu et al. 83 and only “non-burst” data was used in correlation analysis to obtain statistics representative of the large number of small oligomer particles, ignoring rare large ones. Cross-correlation functions of the two detectors were calculated and fitted with a standard model for 3D diffusion in with rapid dye blinking: With amplitude 𝐺0, diffusion time 𝜏d, point spread function aspect ratio S, and blinking parameters FB and 𝜏B. Custom software was written in Python for burst removal, correlation, and fitting, based on tttrlib 0.0.19 (https://github.com/Fluorescence-Tools/tttrlib). The software is in ongoing development, and available upon request.

Image analysis

Electron microscopy images were processed and analyzed using IMOD and ImageJ. The dimensions of the toroidal structures, FtsZ bundles and distances between filaments were manually measured using ImageJ and IMOD. Distances were plotted as histograms using Origin (OriginPro, Version 2019b. OriginLab Corporation, Northampton, MA, USA.). For each toroid analyzed (n = 67) the inner and outer diameter were measured collecting the major and minor distance in each case. The circularity of the toroid was the result of the division between the minor diameter divided by major diameter for each toroid. The height of the FtsZ toroid was manually measured from the tomograms, collecting 4 measurements per toroid to assure a correct representation of the size (n = 17). For the spacing between filaments, the space in between the filaments was measured manually in the XY plane by using IMOD. Each measurement represents only one FtsZ double filament or one FtsZ-ZapD double filament, however, the same bundle could be measured more than once as they are composed for multiple filaments. The measurements were collected from >3 independent samples. The distances between ZapDs connecting two FtsZ filaments were measured following the same methodology. The mean value and standard deviation of different datasets were calculated and added to the figures together with the n used for each case.